---
id: 5d8a4cfbe6b6180ed9a1c9e0
title: Parte 3
challengeType: 0
dashedName: part-3
---

# --description--

No cabeçalho, adicione um `title` com o conteúdo `D3 Dashboard`.

# --hints--

test-text

```js
assert(
  /<head\s*>\s*<title\s*>D3 Dashboard<\/title\s*>\s*<\/head\s*>/g.test(code)
);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html>
  <head>

  
  </head>

  <body>
  </body>
</html>
```

# --solutions--

```html
<!DOCTYPE html>
<html>
  <head>
    <title>D3 Dashboard</title>
  </head>

  <body>
  </body>
</html>
```
