---
id: 5d8a4cfbe6b6180ed9a1c9de
title: Parte 1
challengeType: 0
dashedName: part-1
---

# --description--

Bem vindo ao painel de controle do projeto. Você usará a biblioteca para visualização de dados do Javascript, D3, para construir a visualização dos seguidores da sua rede social. Consistirá de um gráfico de linha, um gráfico de pizza e uma legenda

Primeiro, você precisa criar um arrquivo HTML. Comece adicionando a declaração `<!DOCTYPE html>` ao topo do arquivo para informar ao navegador que tipo de documento está sendo lido.

# --hints--

Texto de teste

```js
assert(/<!DOCTYPE\s+html\s*>/gi.test(code));
```

# --seed--

## --seed-contents--

```html
```

# --solutions--

```html
<!DOCTYPE html>
```
