---
id: 5d8a4cfbe6b6180ed9a1c9de
title: Parte 1
challengeType: 0
dashedName: part-1
---

# --description--

Bem-vindo ao projeto do painel! Você usará a biblioteca de visualização de dados JavaScript, D3, para construir uma visualização de seus seguidores de mídia social. Ele consistirá em um gráfico de linha, um gráfico de pizza e uma legenda.

Primeiro, você precisa criar um arquivo HTML. Inicie adicionando a declaração `<!DOCTYPE html>` ao topo do arquivo para chamar ao navegador que tipo de documento ele está lendo.

# --hints--

test-text

```js
assert(/<!DOCTYPE\s+html\s*>/gi.test(code));
```

# --seed--

## --seed-contents--

```html
```

# --solutions--

```html
<!DOCTYPE html>
```
