---
id: 5d822fd413a79914d39e98cc
title: Parte 4
challengeType: 0
dashedName: part-4
---

# --description--

Aninhe as tags de abertura e fechamento `title` e` style` na área de cabeçalho e dê ao seu projeto o título de `freeCodeCamp Skyline Project`.Todos os estilos que você for solicitado a adicionar durante este projeto devem ir nesta área de estilo.

# --hints--

texto de teste

```js
assert(
  code.match(
    /<head\s*>\s*(<style\s*>\s*<\/style\s*>\s*<title\s*>freeCodeCamp Skyline Project<\/title\s*>|<title\s*>freeCodeCamp Skyline Project<\/title\s*>\s*<style\s*>\s*<\/style>)\s*<\/head\s*>/g
  )
);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html>
  <head> </head>

  <body></body>
</html>
```

# --solutions--

```html
<!DOCTYPE html>
<html>
  <head>
    <title>freeCodeCamp Skyline Project</title>
    <style></style>
  </head>

  <body></body>
</html>
```
