---
id: 5d822fd413a79914d39e98c9
title: Parte 1
challengeType: 0
dashedName: part-1
---

# --description--

Bem-vindo ao projeto de horizonte de Variáveis CSS! Comece adicionando a declaração `! DOCTYPE html` no topo do documento para que o navegador saiba que tipo de documento está lendo.

# --hints--

test-text

```js
assert(code.match(/<!DOCTYPE\s+html\s*>/gi));
```

# --seed--

## --seed-contents--

```html

```

# --solutions--

```html
<!DOCTYPE html>
```
