---
id: 5d822fd413a79914d39e98ca
title: Parte 2
challengeType: 0
dashedName: part-2
---

# --description--

Adicione as tags `html` de abertura e fechamento abaixo do doctype para que você tenha um lugar para começar a colocar algum código.

# --hints--

texto de teste

```js
assert(code.match(/<!DOCTYPE\s+html\s*>\s*<html\s*>\s*<\/html\s*>/gi));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
```

# --solutions--

```html
<!DOCTYPE html>
<html></html>
```
